```python
!pip install multiset
```

```python
from multiset import Multiset
from IPython.display import display
```

#### What is the difference between these two syntaxes?
- `list((1,2,3,4))`
- `[1,2,3,4]`

<!-- grade_id: cell-911bf575d845675f, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

#### Write a list comprehension

In the cell below write a function that uses a list comprehension that takes a numerical list and replaces all of the even numbers with 0. 

You will receive:

- five point for passing the base case ie. writing a function that does the above
- five points if you write error handling should the function be handed a non-numeric list
   - to pass the test, if the function receives a list that contains non-numeric elements, simply return the string `"Can not transform list with non-numeric elements"`.

Your function should look like this:

    def change_evens_to_zeros(lst):
        # do some things
        return modified_lst

<!-- grade_id: cell-a5e86b16fd350b60, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-960a0e875c1c29d3, points: 5, grade: true, solution: false -->
```python
assert change_evens_to_zeros(range(10)) == [0, 1, 0, 3, 0, 5, 0, 7, 0, 9]
```

<!-- grade_id: cell-953a0a79c7aa7453, points: 5, grade: true, solution: false -->
```python
assert change_evens_to_zeros([1,2,3,4,"five"]) == "Can not transform list with non-numeric elements"
```

#### FREE RESPONSE

Why do you think it would be useful to have a compound data type that requires all elements to have the same primitive data type?

<!-- grade_id: cell-ebcadae30f3b2472, points: 7, grade: true, solution: true -->
```python
# your code here
raise NotImplementedError
```

# Lists

## Primitive Data Types 

Primitive data types are the most basic datatypes available to a programming language. The most common primitive datatypes in Python are:

```python
display(type(1))
display(type(1.))
display(type(True))
display(type(None))
display(type(1+1.j))
display(type('None'))
display(type("No'n'e"))
display(type("""

mult-line string

"""))
```

## Compound Data Types

A compound data type is a data type that is made up of one or more primitive data types.

Let's consider three compound data types:
- `set`
- `Multiset`
- `list`

The **length** of an instance of a compound data type is the number of elements in that instance.

The length of a `set` is the number of *unique* elements in that set.

```python
display(len(set((1,1,2,3))))
display(len(set((1,2,3))))
```

The length of a `Multiset` is the number of elements in the `Multiset`.

```python
display(len(Multiset((1,1,2,3))))
display(len(Multiset((1,2,3))))
```

The length of a `list` is the number of elements in the `list`.

```python
display(len(list((1,1,2,3))))
display(len(list((1,2,3))))
```

#### So what is the difference between a `Multiset` and a `list`?

In addition to having a length, a list has an **order**. 

Neither a `set` nor a `multiset` have order.

```python
set((1,1,2,3)) == set((1,2,1,3))
```

```python
Multiset((1,1,2,3)) == Multiset((1,2,1,3))
```

```python
list((1,1,2,3)) == list((1,2,1,3))
```

Two lists are considered equal if and only if they have the same elements in the same order:

```python
list((1,2,3,4)) == [1,2,3,4]
```

## Mathematical Definition of a List

**Definition**: *list*, *length*

A list of length $n$ is an ordered collection of $n$ items separated by commas and surrounded by brackets. 

We might be tempted to think of a list as a list of numbers

```python
a_list = [1,2,3,4]
a_list
```

But in Python, a list is not restricted to only numbers, and is not even restricted to everything being the same primitive type

```python
another_list = ["one", 2., "III", 4]
another_list
```

Both of these are lists and have a length.

```python
print("The length of a_list is %s." % len(a_list))
print("The length of another_list is %s." % len(another_list))
```

It is worth noting that Python lists are indexed from 0 so that the first element is found using the index 0

```python
a_list[0]
```

```python
another_list[2]
```

## List Comprehensions

In this next section, we will use a list comprehension as a tool in our programming. A list comprehension is used to transform compound objects. They take this form

    [do_something_to(placeholder_var) for placeholder_var in compound_object] 

For example, we might trivially wish to convert the list `[1,2,3,4]` to the list `[2,3,4,5]`. This could be done using a list comprehension as

    transformed_list = [v+1 for v in [1,2,3,4]]

We should think of list comprehensions as **vectorized operations** in that we are applying some operation to the entire list, not just one element in the list.

### ADVANCED: `if` and `if-else` statements in list comprehensions

It is trivial to use `if` syntax to filter values from a list comprehension.

    [v for v in compound_object if condition_is_true]

```python
[v for v in [1,2,3,4] if v > 2]
```

Notice that this list only returns the values greater than 2. 

We could even modify these values.

```python
[v**2 for v in [1,2,3,4] if v > 2]
```

`if-else` syntax is a bit trickier, but can actually be used to conditionallly modify a list using a list comprehension.

    [v if condition_is_true else other_val for v in compound_object]

```python
[a if a > 2 else 0 for a in (1,2,3)]
```

```python
[a**2 if a > 2 else a - 4 for a in [1,2,3,4]] 
```

## Homogenous Lists

Mathematically, we typically require that lists contain values all of the same primitive data type. What if we want to define a list where this is true?


Here, we define a new class `hlist` that inherits from the original `list` class. It initializes the `list` using the same method 

    list.__init__(self, *args)
    
but then checks the types of everything in the list

    types = set([type(v) for v in self])
    
by creating a set of the datatypes of the elements of the list. If all of the data types are the same, the `set` will have a length of 1! If the set of data types does not have a length of 1, this means that we have more than one data type in the list. In this case, we `raise` the error `TypeError("All elements of the list must have the same type.")`.

```python
class SomeClass(object):
    def __init__(self, arg):
        self.arg = arg
```

```python
my_instance_of_some_class = SomeClass(2)
```

```python
my_instance_of_some_class.arg
```

```python
class hlist(list):
    def __init__(self, *args):
        list.__init__(self, *args)

        types = set([type(v) for v in self])
        if len(types) != 1:
            raise TypeError("All elements of the list must have the same type.")
```

##### Let's test it out

```python
try:
    this_list = hlist((1,2,3,4))
    print(this_list)
except TypeError as e:
    print(e)
```

```python
try:
    a_bad_list = hlist((1,2,"three"))
    print(a_bad_list)
except TypeError as e:
    print(e)
```

```python
try:
    a_string_list = hlist(("one", "II", "three"))
    print(a_string_list)
except TypeError as e:
    print(e)
```

Are mixed number types ok, though?

```python
try:
    mixed_number_list = hlist((1., 2, 3))
    print(mixed_number_list)
except TypeError as foo:
    print(foo)
```

Mixed number typs are causing the `hlist` to `raise` an error. 

```python
set([type(v) for v in (1.,1.,3)])
```

We can see that the set of data types includes two types, `int` and `float`.

Let's modify the homogenous list class to allow for multiple primitive types *if* they are all numeric.

We will take the following as numeric types:

- `int`
- `float`
- `complex`

Here, we form a set of the data types in our list, but before we do this, we use an `if`-`else` list comprehension to change label for all numeric types to the string `'numeric'`. Consider this

```python
some_list = [1.,2,4,1+0j,'foo']
```

```python
types = [type(v) for v in some_list]
types
```

```python
{'a':1, 'b':1}
```

```python
class SomeClass():
    
```

```python
def some_func(*args):
    return list(args)
```

```
for i in 
```


```python
some_func('a',[1,2,3],{'a':1},{1,2,3})
```

```python
some_func(1,2,3,'a','n')
```

```python
from sklearn.tree import DecisionTreeClassifier
```

```python
tree = DecisionTreeClassifier()
```

```python
{int, float, complex} == set((int, float, complex))
```

```python
numeric_types = {int, float, complex}
```

```python
types = ['numeric' if t in numeric_types else t 
         for t in types]
types
```

```python
set(types)
```

```python
class hlist(list):
    def __init__(self, *args):
        list.__init__(self, *args)

        numeric_types = {int, float, complex}
        
        types = [type(v) for v in self]
        types = ['numeric' if t in numeric_types 
                       else t
                 for t in types]
        unique_types = set(types)
        
        if len(unique_types) != 1:
            raise TypeError("All elements of the list must have the same type.")
```

```python
mixed_number_list = hlist((1., 2, 3))
```

```python
mixed_number_list
```

```python
import numpy as np 
```

```python
np.array(mixed_number_list, dtype=int)
```

```python
np.array((1E-100, 0, 0+1j, False), dtype=bool)
```

```python

```