# Problem Set 1 - The Iris Dataset

```python
import matplotlib.pyplot as plt
import pandas as pd
import seaborn as sns
from IPython.display import display

%matplotlib inline
```

The Iris dataset is available from the UCI Machine Learning Repository at https://archive.ics.uci.edu/ml/machine-learning-databases/iris/iris.data.

The features are given in the following order:

    sepal length
    sepal width
    petal length
    petal width
    flower class 

Please use the following as the column names in your dataframe:

```python
columns = [
    'sepal_length',
    'sepal_width',
    'petal_length',
    'petal_width',
    'flower_class'
]
```

## Load the Iris Dataset

Read the Iris data into a Pandas DataFrame using the `pd.read_csv()` function. Load the data using the URL to the data set at the UCI Machine Learning Repository.

<!-- grade_id: cell-c12569356effd436, points: 0, grade: false, solution: true -->
```python
def load_dataframe():
    """
    Write a function that returns the iris dataset loaded in a dataframe. 
    Make sure that the columns are properly named using the variable
    columns defined in the previous cell and that header rows
    are properly handled.
    """
    # Write your code here.
    
    return iris_df
```

Your `load_dataframe` function should pass all of these tests. If all tests pass, you will see the text `"All tests pass."` following the execution of the cell.

<!-- grade_id: cell-3332c8262d601bde, points: 5, grade: true, solution: false -->
```python
iris_df = load_dataframe()
assert type(iris_df) == pd.DataFrame
assert iris_df.shape == (150,5)
assert list(iris_df.columns) == columns
assert list(iris_df.iloc[4]) == [5.0,3.6,1.4,0.2, "Iris-setosa"]
print("All tests pass.")
```

Display the head of your DataFrame to verify that the data was properly loaded.

```python
iris_df.head()
```

Display the shape of your DataFrame to verify that the DataFrame has the correct shape 

```python
iris_df.shape
```

## Describe the DataFrame

Write a function that returns the description of a DataFrame.

<!-- grade_id: cell-1f7ab55565ad40c2, points: 0, grade: false, solution: true -->
```python
def dataframe_description(dataframe):
    """
    Write a function that returns the description of a dataframe.
    """
    # Write your code here.
    
    return description
```

Your `dataframe_description` function should pass these tests. If all tests pass you should see the text `"All tests pass."` and the description of the Iris Dataframe.

<!-- grade_id: cell-a60bed8c3a336cca, points: 3, grade: true, solution: false -->
```python
assert dataframe_description(iris_df).shape == (8,4)
assert list(dataframe_description(iris_df).index) == [
    'count','mean','std','min',
    '25%','50%','75%','max'
]
assert list(dataframe_description(iris_df).loc['min']) == [4.3,2.0,1.0,0.1]
print("All tests pass.")
dataframe_description(iris_df)
```

## Use `sns.boxplot()` to Display Box Plots of the Features

```python
sns.boxplot(iris_df)
```

## Use `sns.distplot` to display Distribution Plots of the Features

```python
_, ax = plt.subplots(1,4, figsize=(20,5))
sns.distplot(iris_df.sepal_length, ax=ax[0])
sns.distplot(iris_df.sepal_width, ax=ax[1])
sns.distplot(iris_df.petal_length, ax=ax[2])
sns.distplot(iris_df.petal_width, ax=ax[3])
```

## True/False

1. Sepal Width is qualitative.
2. Sepal Length is quantitative.
3. Petal Width is quantitative.
4. Petal Length is qualitative.
5. Flower Class is qualitative.
6. If predicting Flower Class, this is a classification problem. 

<!-- grade_id: cell-631a6c6c400de3ef, points: 0, grade: false, solution: true -->
```python
def true_false_1():
    # Write your code here.
    
    return answer

def true_false_2():
    # Write your code here.
    
    return answer

def true_false_3():
    # Write your code here.
    
    return answer

def true_false_4():
    # Write your code here.
    
    return answer

def true_false_5():
    # Write your code here.
    
    return answer

def true_false_6():
    # Write your code here.
    
    return answer
```

<!-- grade_id: cell-827a886b6e157576, points: 3, grade: true, solution: false -->
```python
# These tests are Hidden.

```

## Numerical Response

1. What is the range of petal width?
2. What is the mean of sepal length?
3. What is the median of petal length?
4. What is the standard deviation of sepal width?

<!-- grade_id: cell-b38f030fd65ec9d0, points: 0, grade: false, solution: true -->
```python
def numerical_1():
    """Express your answer as a tuple of two floating-point numbers 
    each with one number after the decimal."""
    # Write your code here.
    
    return answer

def numerical_2():
    """Express your answer a floating-point number 
    with two numbers after the decimal."""
    # Write your code here.
    
    return answer

def numerical_3():
    """Express your answer a floating-point number 
    with two numbers after the decimal."""
    # Write your code here.
    
    return answer

def numerical_4():
    """Express your answer a floating-point number 
    with two numbers after the decimal."""
    # Write your code here.
    
    return answer
```

<!-- grade_id: cell-7df840b1890dbc9c, points: 3, grade: true, solution: false -->
```python
# These tests are Hidden.

```